<a href="https://colab.research.google.com/github/Adnaan-khan/weekly-assignments-/blob/main/week06_assignment_(3).ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def allergy_count(json_file):
  """Counts the number of allergies in the specified JSON file.

    Parameter:
        json_file (str): The path to the JSON file.

    Returns:
        int: The total number of entries in the 'entry' list.
        example:
        >>> allergy_count(ALLERGIES_FILE)
        4
"""

  with open(json_file) as f: ### f = file
        data = json.load(f)
  return len(data["entry"])
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
import json
"""
Count the unique patients in the given JSON file that contains allergy information.

Parameters:
    json_file (str): The path to the JSON file.

Returns:
    int: The number of distinct patients found in the 'entry' list.
    example:
    >>> patient_count(ALLERGIES_FILE)
    2

"""

ALLERGIES_FILE="allergies.json"
def patient_count(json_file):
    with open(json_file) as f:  ### (f = File)
        data = json.load(f)
        ### Patient Count
        unique_patients = {entry["resource"]["patient"]["display"]
                           for entry in data ["entry"]
                           if entry.get("resource") and entry["resource"].get("patient")}
        return len(unique_patients)
        count = patient_count(ALLERGIES_FILE)
        print(count)





```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
import json

ALLERGIES_FILE = "allergies.json"

def allergy_per_patient(json_file):
  """
    Counts the number of allergies for each patient in the specified JSON file.

    Parameters:
        json_file (str): The path to the JSON file.

    Returns:
        Patient allergy count : A dictionary where keys are patient names and values are the number of allergies they have.
        example:
    >>> allergy_per_patient(ALLERGIES_FILE)
    {'Paul Boal': 1, 'Jason Argonaut': 3}
    """
  with open(json_file) as f:  ### (f = File)
        data = json.load(f)
        entries = data['entry']

        # Gathering all patient displays
        patient_display = [entry['resource']['patient']['display'] for entry in entries]

        # Calculating allergies per patient by using a dictionary
        patient_allergy_count = {patient: patient_display.count(patient) for patient in set(patient_display)}

        return patient_allergy_count

        count = allergy_per_patient(ALLERGIES_FILE)
        print(count)


```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json
def allergy_list(json_file):
  """
    generate a list of allergies and their reactions for each patient from the specified JSON file.

    Parameters:
        json_file (str): The path to the JSON FILE

    Returns:
        list: A list of allergy, where each inner list contains the patient name, allergy substance,
              and reaction manifestation for each allergy entry.

    Example:
    >>> allergy_list("allergies.json")
    ['Jason Argonaut', 'PENICILLIN G', 'Hives'],

    >>> allergy_list("allergies.json")
    ['Paul Boal', 'PENICILLIN G', 'Bruising'],

    >>> allergy_list("allergies.json")
    ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],

    >>> allergy_list("allergies.json")
    ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']


    """
  with open(json_file) as f:  ### (f = File)
        data = json.load(f)
        entries = data["entry"]

        allergy_list = []
        for entry in entries:
            allergy_list.append([
                entry.get('resource', {}).get('patient', {}).get('display'),
                entry.get('resource', {}).get('substance', {}).get('text'),
                entry.get('resource', {}).get('reaction', [{}])[0].get('manifestation', [{}])[0].get('text')
            ])
        allergy_list.sort()
        return allergy_list






```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]


```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

def allergy_reaction(json_file, patient, diagnosis):
  """
     Retrieves the allergic reaction for a specific patient and diagnosis from the JSON file.

     Parameters:
     json_file (str): The path to the JSON file.
     patient (str): The name of the patient.
     diagnosis (str): The diagnosis or medication related to allergic reaction.

     Returns:
     str or None: The allergic reaction if found; otherwise, returns None.

     example:
     >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
     'Hives'
     >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS')
     'Itching'
     >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY')
     'Anaphylaxis'
     >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN')
     >>> allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')
     'Bruising'
  """
  with open(json_file) as f:  ### (f = File)
      data = json.load(f)
      entries = data["entry"]

      allergy_reaction = {}

      for entry in entries:
        patient_name = entry.get('resource').get('patient').get('display');
        medicine = entry.get('resource').get('substance').get('text');
        reaction = entry.get('resource').get('reaction')[0].get('manifestation')[0].get('text');

        allergy_reaction.update({patient_name + medicine: reaction})

      key = patient+diagnosis


  return allergy_reaction.get(key,None)





```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

